<problem display_name="Example 4" linenum="402" tex_filename="./block3/basis.tex" showanswer="finished" attempts="3" rerandomize="per_student">
  <text>
    <p>
Let [mathjaxinline]v_1=\left[\begin{array}{c} 1 \\ 2 \\ 3 \end{array} \right][/mathjaxinline] and [mathjaxinline]v_2 = \left[\begin{array}{c} 3 \\ 0 \\ 1 \end{array} \right][/mathjaxinline]. Then [mathjaxinline]\mathcal{B} = \{ v_1; v_2\}[/mathjaxinline] is a basis for some subspace [mathjaxinline]W[/mathjaxinline] in [mathjaxinline]\mathbb R^3[/mathjaxinline]. </p>
    <p>
What vector [mathjaxinline]w[/mathjaxinline] has the property that [mathjaxinline][w]_{\mathcal{B}} = \left[\begin{array}{c} 1 \\ 2 \end{array} \right][/mathjaxinline]? </p>
    <script linenum="2" filename="./vectorentry2.tex" type="text/python" system_path="python_lib">
def VectorEntry(expect, ans):
	import ast
	import numpy as np 
  	atol = 0.01
	list_expect = ast.literal_eval(expect)
	vec_expect = np.matrix(list_expect)
  	ret = {"ok":False}
	try:
  		# input format [[1],[2],[3]]
		list_ans = ast.literal_eval(ans)
		vec_ans = np.matrix(list_ans)
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,atol,1e-08):
  			ret['ok'] = True
  		else:
  		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong'   			
	#except SyntaxError:
		#ret['msg'] = 'Wrong input format'
	except SyntaxError:
  		# input format &amp;lt;1,2,3&amp;gt;
		list_ans = ans.replace('&amp;lt;', '[').replace('&amp;gt;', ']')
		list_ans = ast.literal_eval(list_ans)
		vec_ans = np.transpose(np.matrix(list_ans))
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
    		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong' 
  	except:
  		ret['msg'] = 'Wrong input format'
  	return ret 
</script>
    <p>
      <customresponse cfn="VectorEntry" expect="[[7],[2],[5]]">
        <textline correct_answer="[[7],[2],[5]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 If [mathjaxinline][w]_{\mathcal{B}} = \left[\begin{array}{c} 1 \\ 2 \end{array} \right][/mathjaxinline], then [mathjaxinline]w = 1v_1 + 2v_2.[/mathjaxinline] Thus we compute [mathjaxinline]q[/mathjaxinline] to be [mathjaxinline]\left[\begin{array}{c} 7 \\ 2 \\ 5 \end{array} \right].[/mathjaxinline]  </p>
        </font>
      </solution>
    </p>
  </text>
</problem>
